Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Final Project - Connect 4
**COMPSCI 101, Duke Kunshan University**

In this project you will create a text-based version of Connect Four and develop a rudimentary AI player that can look at all possible boards up to a certain number of moves.

Connect Four (also known as Captain's Mistress, Four Up, Plot Four, Find Four, Four in a Row, Four in a Line, Drop Four, and Gravitrips (in Soviet Union)) is a two-player connection game in which the players first choose a color and then take turns dropping one colored disc from the top into a seven-column, six-row vertically suspended grid.

More information can be found here: https://en.wikipedia.org/wiki/Connect_Four

# BOARD

Write a class called `Board` that will represent the Connect Four 6 row x 7 column vertical grid using a 2-D Python List. Pieces that are played will be represented by the strings: 'X' and 'O'. **Note that** these two strings 'X' and 'O' are called `marker`.

1. Write the `__init__` method that takes two integer inputs, `cols` and `rows`, with default values 7 and 6 respectively. This should set up an attribute `data` that is the 2-D Python List along with attributes to store the number of columns and rows.


2. Write the `__str__` method that returns a string that represents the vertical grid including column numbers at the bottom starting from 0 to cols-1. It should look something like below. **Note that** the first/top row is `data[0]`.

<pre>
| | | | | | | |
| | | | | | | |
| | | | | | | |
| | | | | | | |
| | | | | | | |
| | | | | | | |
---------------
 0 1 2 3 4 5 6 
</pre>


3. Write a method called `is_valid_move` that takes an integer `col` as input and returns `True` if a move can be made on that column and `False` otherwise. Be sure to check that `col` is between 0 and col-1 and there are empty slots in the column of choice.


4. Write a method called `add_marker` that takes two inputs, an integer `col` and a string `marker`, that adds the `marker` to the specified column. **Note that** you need to find the first empty space in the `col` and then put the `marker` there.


5. Write a method called `del_marker` that takes an integer `col` as input and deletes the top marker in the specified column.


6. Write a method called `reset_board` that clears the board by setting all the locations in the board to ' '.


7. Write a method called `set_board` that takes a string `cols` and a string `marker` that sets up the board according to the columns represented in `cols`. If `cols` is '0123456' this means the first move is at column 0 with the specified marker, then the second move is at column 1 with the opposite marker, then the third move is at column 2 with the opposite marker, etc. This provides an easy way to set up the board and run tests. For example, the method call `set_board('01234561263521','X')` will produce a board like below. **Note that** you may want to call the method `add_marker` for this task.

<pre>
| | | | | | | |
| | | | | | | |
| | | | | | | |
| |O|X| | | | |
| |O|X|X| |O|O|
|X|O|X|O|X|O|X|
---------------
 0 1 2 3 4 5 6
</pre>

8. Write a method called `is_full` that returns `True` if there are no more moves left to play and `False` otherwise.


9. Write a method called `is_winner` that takes a string `marker` as input and returns `True` if `marker` wins and `False` otherwise. **Note that** a win can occur horizontally, vertically, or diagonally, so you must check all these directions.  

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
class Board():
    def __init__(self, cols=7, rows=6):
        # YOUR CODE HERE
        self.cols = cols
        self.rows = rows
        self.data = [ ['| ' for i in range(self.cols) ] for j in range(self.rows) ]
        
    def __str__(self):
        # YOUR CODE HERE
        boardstring = ''
        for r in range(self.rows):
            for c in range(self.cols):
                boardstring += self.data[r][c]
            boardstring += '|\n' 
        boardstring = boardstring + '-'*15 + '\n'
        numstr = ''
        for i in range(7):
            numstr += ' '+ str(i)
        return boardstring + numstr + ' \n'

    def is_valid_move(self, col):
        # YOUR CODE HERE
        if col > self.cols-1 or col < 0:
            return False
        else:
            if self.data[0][col] == '| ':
                return True
            else:
                return False
       
    
    def add_marker(self, col, marker):
        # YOUR CODE HERE
        for i in range(self.rows-1,-1,-1):
            if self.data[i][col] == '| ':
                self.data[i][col] = '|' + marker
                break
                
            
    def del_marker(self, col):
        # YOUR CODE HERE
        for i in range(self.rows):
            if self.data[i][col] == '|X' or self.data[i][col] == '|O':
                self.data[i][col] = '| '
                break
    
    def reset_board(self):
        # YOUR CODE HERE
        for i in range(self.rows):
            for j in range(self.cols):
                if self.data[i][j] == '|X' or self.data[i][j] == '|O':
                    self.data[i][j] = '| '

    def set_board(self, cols, marker):
        # YOUR CODE HERE
        self.reset_board()
        for i in range(len(cols)):
            numcol = int(cols[i])
            if marker == 'X':
                if i % 2 == 0:
                    self.add_marker(numcol, 'X')
                else:
                    self.add_marker(numcol, 'O')
            else:
                if i % 2 == 0:
                    self.add_marker(numcol, 'O')
                else:
                    self.add_marker(numcol, 'X')
                
    def is_full(self):
        # YOUR CODE HERE
        for i in range(self.rows):
            for j in range(self.cols):
                if self.data[i][j]== '| ':
                    return False
        return True
    
    def is_winner(self, marker):
        # YOUR CODE HERE
        check = False
        if marker == 'X':
            for r in range(self.rows):
                for c in range(self.cols-3):
                    if self.data[r][c] == '|X' and self.data[r][c+1] == '|X' and self.data[r][c+2] == '|X' and self.data[r][c+3] == '|X':
                        check = True
            for r in range(self.rows-3):
                for c in range(self.cols):
                    if self.data[r][c] == '|X' and self.data[r+1][c] == '|X' and self.data[r+2][c] == '|X' and self.data[r+3][c] == '|X':
                        check = True
            for r in range(self.rows-3):
                for c in range(self.cols-3): 
                    if self.data[r][c] == '|X' and self.data[r+1][c+1] == '|X' and self.data[r+2][c+2] == '|X' and self.data[r+3][c+3] == '|X':
                        check = True
            for r in range(3,self.rows):
                for c in range(self.cols-3): 
                    if self.data[r][c] == '|X' and self.data[r-1][c+1] == '|X' and self.data[r-2][c+2] == '|X' and self.data[r-3][c+3] == '|X':
                        check = True
        else:
            for r in range(self.rows):
                for c in range(self.cols-3):
                    if self.data[r][c] == '|O' and self.data[r][c+1] == '|O' and self.data[r][c+2] == '|O' and self.data[r][c+3] == '|O':
                        check = True
            for r in range(self.rows-3):
                for c in range(self.cols):
                    if self.data[r][c] == '|O' and self.data[r+1][c] == '|O' and self.data[r+2][c] == '|O' and self.data[r+3][c] == '|O':
                        check = True
            for r in range(self.rows-3):
                for c in range(self.cols-3): 
                    if self.data[r][c] == '|O' and self.data[r+1][c+1] == '|O' and self.data[r+2][c+2] == '|O' and self.data[r+3][c+3] == '|O':
                        check = True
            for r in range(3,self.rows):
                for c in range(self.cols-3): 
                    if self.data[r][c] == '|O' and self.data[r-1][c+1] == '|O' and self.data[r-2][c+2] == '|O' and self.data[r-3][c+3] == '|O':
                        check = True
        return check
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 3, grade: true, solution: false -->
```python
board = Board()

assert str(board) == '| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n---------------\n 0 1 2 3 4 5 6 \n'
assert board.is_valid_move(0) == True
assert board.is_valid_move(4) == True
assert board.is_valid_move(8) == False

board.add_marker(0,'X')
board.add_marker(1,'O')
board.add_marker(1,'X')
board.add_marker(1,'O')

assert str(board) == '| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| |O| | | | | |\n| |X| | | | | |\n|X|O| | | | | |\n---------------\n 0 1 2 3 4 5 6 \n'

board.del_marker(1)

assert str(board) == '| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| |X| | | | | |\n|X|O| | | | | |\n---------------\n 0 1 2 3 4 5 6 \n'

board.reset_board()

assert str(board) == '| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n---------------\n 0 1 2 3 4 5 6 \n'

board.set_board('0123456243425','X')

assert str(board) == '| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| | |O| |X| | |\n| | |O|O|X|X| |\n|X|O|X|O|X|O|X|\n---------------\n 0 1 2 3 4 5 6 \n'

assert board.is_full() == False
assert board.is_winner('X') == False
assert board.is_winner('O') == False

board.set_board('012345601234560123456012345601234560123456','X')

assert board.is_full() == True
assert board.is_winner('X') == True
assert board.is_winner('O') == True
```

# Player

Write a class called `Player` that models a human or an AI player that is determined by the `kind` attribute. Each player has a `marker` attribute that is either 'X' or 'O', and an `opp_marker` attribute of opponent that is the other possible marker. For example, if `marker` is 'X', `opp_marker` will be 'O'.

If the player is an AI player, then it will also include a `strategy` attribute that is either 'LEFT', 'RIGHT' or 'RAND', and a `level` attribute that determines the number of levels of future moves it will consider when making a move. For example, if `level` is 1, AI player will think one step further, "If I make this move, where may my opponent move?" In other words, find a move with highest score that is explained later.

1. Write the `__init__` method that takes up to four string inputs, `marker` ('X' or 'O'), `kind` ('HUMAN' or 'AI'), `strategy` ('LEFT','RIGHT','RAND') and `level` (any number > 0). The `strategy` and `level` attributes are optional and default to 'RAND' and 0 respectively. Your class should keep track of the opponent's marker as well.


2. Write the `__str__` method that returns a string that represents information about the player that should look something like one of lines below, depending on the `kind` attribute:

    <pre>
    HUMAN (X)
    AI (O): LEFT, lvl: 2
    </pre>

    It tells us whether the player is a human or an AI, what marker it is using and if it is an AI, what strategy it is using and what level it is using its strategy to.


3. Write a method called `human_move` that takes a `Board` as a parameter. The method asks input from the human player (using the Python built-in method `input`). The input is a position in the board, i.e, a column. It must be a valid move. You may use the method that you defined in `Board` to verify. Keep asking for an input until a valid move is given.


4. Write a method called `choose_move` that takes a list of `scores` (that are indexed by moves) and, based on the `strategy` returns the move with the best score. The input below might run as follows:

    <pre>
    [0, 0.5, 0.5, 0, 0.5, 0.5, 0]
    </pre>

    Each index of `scores` represents the move and the value at that index represents the score for that move. You can see that the values at indices 1, 2, 4, and 5 are all 0.5 - these can all be considered best moves. The `strategy` determines which of these moves will be played. If the strategy is 'LEFT' then the left-most move is chosen: 1. If the strategy is 'RIGHT' then the right-most move is chosen: 5. If the strategy is 'RAND' then a random move (among 1, 2, 4, and 5) is chosen. When choosing a random move you must use the random.randint() function. 
    
    **Note that** you may want to create another list that holds all best moves. For example, if `scores` = [0, 0.5, 0.5, 0, 0.5, 0.5, 0], a list with best moves will be [1, 2, 4, 5].


5. Write a method called `get_move_scores` that will return a list with scores for each possible move that can be made. A move has a score of 1 if it guarantees a win. A move has a score of 0 if it guarantees a loss. A move has a score of 0.5 if it neither guarantees a win nor guarantees a loss. 

    There are several possibilities for judging the score of a given move:

    - If the move is invalid, then its score is -1.
      
    - If the board is already a winning board, then its score is 1.
      
    - If the board is already a losing board, its score is 0.
      
    - If the 'level' is 0 (meaning there is no assessment of future moves) then its score is 0.5.
   
    - Finally, if the move is valid, the board is neither a winning or losing board, and the 'level' is greater than 0, then you can start evaluating future moves. In order to give a score to this move, you must FIRST make the move, i.e. add your marker to the relevant column, and then check if the move is a winning move. If it is,  the score for this move is 1. Otherwise, you must check what the best move for your opponent will be, given your move. This can be done by creating an opponent player with the same AI strategy and level-1, then recursively using the get_move_scores method. Your move's score can be calculated by substracting the opponent's highest score by 1. Given your move, if the opponent can win, this move will yield 1-1, which is 0. If the opponent cannot win, this will yield 1-0.5, which is 0.5.
    
    **Note that**, after you make your move to check your opponent's best move, you need to delete your move so that the original board is restored!
    
    (1) Consider the following examples when the level is 0 for the marker 'X' (i.e, the player who puts down 'X'). When the level is 0, no impact to the future moves are considered. Therefore, just check whether or not the move is valid:
    
<pre>
| | | | | | | |        | | | | | | | |     
| | | | | | | |        | | | | | | | |
| | | | | | | |        | | | | | | | |
| | | | | | | |   ->   | | | | | | | |   All moves are equally good
| | | | | | | |        | | | | | | | |
| | | | | | | |        | | | | | | | |
---------------        ---------------
 0 1 2 3 4 5 6         .5.5.5.5.5.5.5
  ( level==0 ) 

</pre>

<pre>
|O| | | | | | |        |O| | | | | | |     
|X| | | | | | |        |X| | | | | | |
|O| | | | | | |        |O| | | | | | |
|X| | | | | | |   ->   |X| | | | | | |   Move 0 is invalid: its score is -1
|O| | | | | | |        |O| | | | | | |
|X| | | | | | |        |X| | | | | | |
---------------        ---------------
 0 1 2 3 4 5 6         -1.5.5.5.5.5.5
  ( level==0 ) 
</pre>

<pre>
|O| | | | | | |        |O| | | | | | |     
|X| | | | | | |        |X| | | | | | |
|O| | | | | | |        |O| | | | | | |
|X|O| | | | | |   ->   |X| | | | | | |   Any move is great because
|O|O| | | | | |        |O|O| | | | | |   'X' has already won
|X|X|X|X| | | |        |X|X|X|X| | | |
---------------        ---------------
 0 1 2 3 4 5 6         -1 1 1 1 1 1 1
  ( level 0 ) 
</pre>

(2) Consider the following examples when the level is 1 for the marker 'X'. The player makes a valid move and then check what the best scores are for the opponent:
    
<pre>
| | | | | | | |        | | | | | | | |        | | | | | | | |   Since the opponent is at 
| | | | | | | |        | | | | | | | |        | | | | | | | |   level 0, all moves are
| | | | | | | |        | | | | | | | |        | | | | | | | |   scored the same - 0.5.
| | | | | | | |   ->   | | | | | | | |   ->   | | | | | | | |   
| | | | | | | |        | | | | | | | |        | | | | | | | |
| | | | | | | |        |X| | | | | | |        |X| | | | | | |
---------------        ---------------        ---------------
 0 1 2 3 4 5 6          0 1 2 3 4 5 6         .5.5.5.5.5.5.5
  ( level==1 )            ( level==0 ) for opp
  
| | | | | | | |
| | | | | | | |   So, the move at 0 position (first column) will score:
| | | | | | | |   1 - max(0.5,0.5,0.5,0.5,0.5,0.5,0.5) = 0.5
| | | | | | | |   Then you repeat this step to check all
| | | | | | | |   possible moves.
|X| | | | | | |   
---------------   
.5 ? ? ? ? ? ?
 </pre>

<pre>
|O| | | | | | |        |O| | | | | | |        |O| | | | | | |   Since the opponent is at 
|X| | | | | | |        |X| | | | | | |        |X| | | | | | |   level 0, and it is a losing
|O| | | | | | |        |O| | | | | | |        |O| | | | | | |   board for the opponent,
|X|O| | | | | |   ->   |X|O| | | | | |   ->   |X| | | | | | |   any of its move is scored as a 0.
|O|O| | | | | |        |O|O| | | | | |        |O|O| | | | | |
|X|X|X| | | | |        |X|X|X|X| | | |        |X|X|X|X| | | |
---------------        ---------------        ---------------
 0 1 2 3 4 5 6          0 1 2 3 4 5 6          0 0 0 0 0 0 0
  ( level==1 )            ( level==0 ) - opp
  
|O| | | | | | |
|X| | | | | | |   So, the move at 3 (4th column) will score: 
|O| | | | | | |   1 - max(0,0,0,0,0,0,0) = 1
|X|O| | | | | |   This should be repeated for all
|O|O| | | | | |   possible moves.
|X|X|X|X| | | |   
---------------   
 0 ? ? 1 ? ? ?
</pre>


6. Write a method called `AI_move` that takes a `Board` as input and returns the 'best' move given the `strategy` and `level` of the AI player. This should just be a simple application of the two previous methods: `get_move_scores` and `choose_move`.


7. Finally, write a method called `next_move` that takes a `Board` as input and returns the next move. If `kind` is 'AI' then return the next best AI move. If `kind` is 'HUMAN' then return the next valid human input.

<!-- grade_id: cell-c49e901548cede52, points: 0, grade: false, solution: true -->
```python
import random

class Player():
    def __init__(self, marker, kind, strategy='RAND', level=0):
        # YOUR CODE HERE
        self.marker=marker
        self.kind=kind
        self.strategy=strategy
        self.level=level
        
    def __str__(self):
        # YOUR CODE HERE
        result=""
        if self.kind=="HUMAN":
            result="HUMAN ("+self.marker+")"
        else:
            result="AI ("+self.marker+"): "+self.strategy+", lvl: "+str(self.level)
        return result
    
    def human_move(self, board):
        # YOUR CODE HERE
        while True:
            n=int(input("please enter you move:"))
            result=board.is_valid_move(n)
            if result:
                return result
                
        
    def choose_move(self, scores):
        # YOUR CODE HERE
        L=[]
        move=0
        max1=max(scores)
        for i in range(len(scores)):
            if scores[i]==max1:
                L.append(i)
        if self.strategy=="LEFT":
            move=L[0]
        elif self.strategy=="RIGHT":
            move=L[-1]
        else:
            n=random.randint(0,len(L)-1)
            move=L[n]
        return move
            

    def get_move_scores(self, board):
        # YOUR CODE HERE
        L=[]
        if self.level==0:
            for i in range(7):
                valid=board.is_valid_move(i)
                if valid:
                    if self.marker=="X":
                        if board.is_winner("X"):
                            L.append(1)
                        elif board.is_winner("O"):
                            L.append(-1)
                        else:
                            board.add_marker(i,self.marker)
                            if  board.is_winner("X"):
                                L.append(1)
                            else:
                                L.append(0.5)
                            board.del_marker(i)
                    else:
                        if board.is_winner("O"):
                            L.append(1)
                        elif board.is_winner("X"):
                            L.append(-1)
                        else:
                            board.add_marker(i,self.marker)
                            if  board.is_winner("O"):
                                L.append(1)
                            else:
                                L.append(0.5)
                            board.del_marker(i)
                else:
                    L.append(-1)
        elif self.level>0:
            for j in range(7):
                valid=board.is_valid_move(j)
                if valid:
                    t=[]
                    if self.marker=="X":
                        if board.is_winner("X"):
                            L.append(1)
                            board.del_marker(j)
                            continue
                        elif  board.is_winner("O"):
                            L.append(-1)
                        else: 
                            board.add_marker(j,self.marker)
                            if board.is_winner("X"):
                                L.append(1)
                            else:
                                opp = Player('O','AI',self.strategy,self.level-1)
                                t=opp.get_move_scores(board)
                                L.append(1-max(t))
                            board.del_marker(j)
                    else:
                        if board.is_winner("O"):
                            L.append(1)
                            board.del_marker(j)
                            continue
                        elif  board.is_winner("X"):
                            L.append(-1)
                        else:
                            board.add_marker(j,self.marker)
                            if board.is_winner("O"):
                                L.append(1)
                            else:
                                opp = Player('X','AI',self.strategy,self.level-1)
                                t=opp.get_move_scores(board)
                                L.append(1-max(t))
                            board.del_marker(j)
                else:
                    L.append(-1) 
        #print(L)
        return L
    
    def AI_move(self, board):
        # YOUR CODE HERE
        return self.choose_move(self.get_move_scores(board))

    def next_move(self, board):
        # YOUR CODE HERE
        if self.kind=="HUMAN":
            board.add_marker(self.human_move(board),self.marker)
        else:
            board.add_marker(self.AI_move(board),self.marker)
            
random.seed(0)    
board = Board()
board.set_board('00000011211','X')
#Player('O','AI','RAND',2).get_move_scores(board) == [-1.0, 0.0, 0.0, 0.5, 0.0, 0.0, 0.0]
print(Player('O','AI','RAND',1).get_move_scores(board) )
```

Output:
```
[-1, 0, 0, 0.5, 0, 0, 0]
```

<!-- grade_id: cell-cd0b0c4a5944a662, points: 6, grade: true, solution: false -->
```python
random.seed(0)

scott = Player('X','HUMAN')
james = Player('O','AI','LEFT',2)

assert str(scott) == 'HUMAN (X)'
assert str(james) == 'AI (O): LEFT, lvl: 2'

board = Board()

assert Player('X','AI','LEFT',0).get_move_scores(board) == [0.5, 0.5, 0.5, 0.5, 0.5, 0.5, 0.5]
assert Player('O','AI','RIGHT',0).get_move_scores(board) == [0.5, 0.5, 0.5, 0.5, 0.5, 0.5, 0.5]
assert Player('X','AI','RAND',0).get_move_scores(board) == [0.5, 0.5, 0.5, 0.5, 0.5, 0.5, 0.5]

assert Player('O','AI','LEFT',0).choose_move([0, 0.5, 0.5, 0.5, 0, 0.5, 0]) == 1
assert Player('O','AI','RIGHT',0).choose_move([0, 0.5, 0.5, 0.5, 0, 0, 0]) == 3
assert Player('O','AI','RAND',0).choose_move([0, 0.5, 0.5, 0.5, 0, 0.5, 0]) == 5
assert Player('O','AI','LEFT',0).choose_move([0, 0.5, 1, 0.5, 0, 0.5, 0]) == 2
assert Player('O','AI','RIGHT',0).choose_move([0, 0.5, 1, 0.5, 0, 0.5, 0]) == 2
assert Player('O','AI','RAND',0).choose_move([0, 0.5, 1, 0.5, 0, 0.5, 0]) == 2

# Given the following board, 'X' can win (in 1 move) if it plays move 3
board.set_board('000000112','X')
assert Player('X','AI','LEFT',1).get_move_scores(board) == [-1.0, 0.5, 0.5, 1.0, 0.5, 0.5, 0.5]
assert Player('X','AI','LEFT',1).AI_move(board) == 3

# Given the following board, 'O' will lose (in 2 moves) if it plays any move other than 3
board.set_board('00000011211','X')
assert Player('O','AI','RAND',2).get_move_scores(board) == [-1.0, 0.0, 0.0, 0.5, 0.0, 0.0, 0.0]
assert Player('O','AI','RAND',2).AI_move(board) == 3
```

# Game

Write a class called `Game` that will let two players (either a human or an AI or both) play Connect 4 against each other.

1. Write the `__init__` method that takes `cols`, `rows` (both integers) and `p1`, `p2` (both players) as inputs (parameters), and creates a board based on the dimensions (`cols` and `rows`) given and maintains attributes for each player.


2. Write the `__str__` method that returns the state of the board and information about each player. It should return a string that looks something like this:

<pre>
| | | | | | | |
| | | | | | | |
| | | | | | | |
| | | | | | | |
| | | | | | | |
| | | | | | | |
---------------
 0 1 2 3 4 5 6 

1: AI (X): RAND, lvl: 1
2: HUMAN (O)
</pre>


3. Write a method called `play` that takes an optional input `display` that has a default value of `True`. If `display` is True then each move of the game will be displayed on the board. If it is False then the game will be played without anything output to the screen.

    Playing a game begins by resetting the board (use the `reset_board` method). As long as the board is not full, the game will continue to alternate moves between players. `p1` plays a move, the board is updated and you must check whether or not `p1` wins. If `p1` wins, the game ends. You must also check to see if the board is full. If it is full, the game ends in a draw. Then `p2`  plays a move, you do the similar check (as to `p1`). This cycle gets repeated until the game ends with a winner or a draw.
    
    This method should also keep track of the total number of moves that are made in the game from both players. Let's call this number - `num_moves`. If `p1` wins, this method returns a tuple (1,num_moves). If `p2` wins, this method returns (-1,num_moves), and if the game ends in a draw this method returns (0,num_moves).

<!-- grade_id: cell-3c955624d75fdfcc, points: 0, grade: false, solution: true -->
```python
class Game():
    def __init__(self, cols, rows, p1, p2):
        # YOUR CODE HERE
        self.p1 = p1
        self.p2 = p2
        self.board = Board(cols,rows)
        
    def __str__(self):
        # YOUR CODE HERE
        str1=str(self.board)
        str2=str(self.p1)
        str3=str(self.p2)
        return str1+"\n1: "+str2+"\n2: "+str3+"\n"
    
    def play(self, display=True):
        # YOUR CODE HERE
        self.board.reset_board()
        n=0
        num_moves=0
        while True:
            num_moves+=1
            if n%2==0:
                self.p1.next_move(self.board)
                if display:
                    print(str(self.board))
                if self.board.is_winner(self.p1.marker):
                    return (1,num_moves)
                elif self.board.is_full():
                    return (0,num_moves)   
            else:
                self.p2.next_move(self.board)
                if display:
                    print(str(self.board))
                if self.board.is_winner(self.p2.marker):
                    return (-1,num_moves)
                elif self.board.is_full():
                    return (0,num_moves)
            n=n+1

```

<!-- grade_id: cell-425f7463c0d93639, points: 3, grade: true, solution: false -->
```python
g = Game(7,6,Player('X','AI','RAND',2),Player('O','HUMAN'))
assert str(g) == '| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n| | | | | | | |\n---------------\n 0 1 2 3 4 5 6 \n\n1: AI (X): RAND, lvl: 2\n2: HUMAN (O)\n'

g = Game(7,6,Player('X','AI','RAND',2),Player('O','AI','RAND',1))

wins = 0
for i in range(100):
    if g.play(False)[0] == 1:
        wins += 1
assert wins > 60
```

```python

```